<!--NAVIGATION-->

<a href="https://colab.research.google.com/github/bpesquet/machine-learning-katas/blob/master/notebooks/coding_from_scratch/classification_metrics.ipynb"><img align="left" src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open in Colab" title="Open in Google Colaboratory"></a>


## Instructions

This is a self-correcting notebook generated by [nbgrader](https://github.com/jupyter/nbgrader). 

Fill in any place that says `YOUR CODE HERE` or `YOUR ANSWER HERE`. Run subsequent cells to check your code.

# Code Regression Metrics From Scratch

## Package setup

```python
import numpy as np
```

## Accuracy

### Question

Complete the following functions to compute MAE, MSE, MAPE and R^2.

<!-- grade_id: cell-6a21120c3e66ee57, points: 0, grade: false, solution: true -->
```python
def mae(y_true, y_pred):
    y_err = abs(y_pred - y_true)
    return sum(y_err)/len(y_err)
    # YOUR CODE HERE
```

```python
def mse(y_true, y_pred):
    return np.mean((y_pred - y_true)**2)
    # YOUR CODE HERE
```

```python
def r2(y_true, y_pred):
    return 1 - mse(y_true, y_pred) / mse(y_true, np.mean(y_true))
    # YOUR CODE HERE
```

```python
def mape(y_true, y_pred):
    return np.mean(abs(y_pred - y_true)/y_true)
    # YOUR CODE HERE
```

```python
y_true = np.array([2, 0, 2, 2, 0, 1])
y_pred = np.array([0, 1, 2, 2, 0, 2])

mae_value = mae(y_true, y_pred)
mse_value = mse(y_true, y_pred)
mape_value = mse(y_true, y_pred)
r2_value = r2(y_true, y_pred)
print(f'MAE: {mae_value:.2f}')
print(f'MSE: {mse_value:.2f}')
print(f'MAPE: {mape_value * 100:.2f}%')
print(f'R^2: {r2_value:.2f}')

assert mae_value == 0.67
assert mse_value == 1.00
assert mape_value == 1
assert r2_value == -0.24
```

<!-- grade_id: cell-29d49954b89dc997, points: 1, grade: true, solution: false -->
```python
y_true = np.array([2, 0, 2, 2, 0, 1])
y_pred = np.array([0, 1, 2, 2, 0, 2])

acc = accuracy_score(y_true, y_pred)
print(f'Accuracy: {acc * 100:.2f}%')

assert acc == 0.5
```

```python

```